---
title: 'Verwendung von RStudio'
description: 'Überblick über die Funktionalitäten von RStudio'
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: ea72ec7f53
xp: 50
```

Für welche der nachfolgend angeführten Aufgaben eignet sich RStudio?

`@possible_answers`
1. [Als Editor zum effizienten schreiben von R-Programmen]
2. Als Compiler für die Umwandlung von R-Codes in lauffähige Programme
3. [Als R-Console zur Eingabe von R-Befehlen]
4. [Als Ausgabeanzeige von R-Programmen]

`@hint`
es sind mehrere Antworten möglich

`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```
